% Running in auto input_syntax mode. Trying TPTP
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on type_relatorType_subsumed_to_kind for (31ds)
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for type_relatorType_subsumed_to_kind
% SZS output start Proof for type_relatorType_subsumed_to_kind
fof(f15450,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f646,f658,f674,f702,f746,f834,f846,f975,f1020,f1040,f1108,f1124,f1257,f1285,f1289,f1293,f1327,f1359,f1375,f1794,f1827,f2132,f2138,f2279,f2285,f2306,f6770,f7048,f8452,f8876,f9141,f14502,f14506,f14513,f15385,f15397,f15430,f15440,f15449])).
fof(f15449,plain,(
  ~spl59_52 | spl59_344 | ~spl59_1524),
  inference(avatar_contradiction_clause,[],[f15448])).
fof(f15448,plain,(
  $false | (~spl59_52 | spl59_344 | ~spl59_1524)),
  inference(subsumption_resolution,[],[f15446,f2305])).
fof(f2305,plain,(
  ~intrinsicMoment(sK41(xyz)) | spl59_344),
  inference(avatar_component_clause,[],[f2304])).
fof(f2304,plain,(
  spl59_344 <=> intrinsicMoment(sK41(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_344])])).
fof(f15446,plain,(
  intrinsicMoment(sK41(xyz)) | (~spl59_52 | ~spl59_1524)),
  inference(resolution,[],[f15439,f845])).
fof(f845,plain,(
  ( ! [X0] : (~mode(X0) | intrinsicMoment(X0)) ) | ~spl59_52),
  inference(avatar_component_clause,[],[f844])).
fof(f844,plain,(
  spl59_52 <=> ! [X0] : (intrinsicMoment(X0) | ~mode(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_52])])).
fof(f15439,plain,(
  mode(sK41(xyz)) | ~spl59_1524),
  inference(avatar_component_clause,[],[f15438])).
fof(f15438,plain,(
  spl59_1524 <=> mode(sK41(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1524])])).
fof(f15440,plain,(
  spl59_1524 | ~spl59_27 | ~spl59_1523),
  inference(avatar_split_clause,[],[f15436,f15428,f744,f15438])).
fof(f744,plain,(
  spl59_27 <=> ! [X0] : (~externallyDependentMode(X0) | mode(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_27])])).
fof(f15428,plain,(
  spl59_1523 <=> externallyDependentMode(sK41(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1523])])).
fof(f15436,plain,(
  mode(sK41(xyz)) | (~spl59_27 | ~spl59_1523)),
  inference(resolution,[],[f15429,f745])).
fof(f745,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) ) | ~spl59_27),
  inference(avatar_component_clause,[],[f744])).
fof(f15429,plain,(
  externallyDependentMode(sK41(xyz)) | ~spl59_1523),
  inference(avatar_component_clause,[],[f15428])).
fof(f15430,plain,(
  spl59_1523 | ~spl59_341 | ~spl59_1012 | ~spl59_1522),
  inference(avatar_split_clause,[],[f15426,f15395,f9139,f2277,f15428])).
fof(f2277,plain,(
  spl59_341 <=> relator(sK41(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_341])])).
fof(f9139,plain,(
  spl59_1012 <=> ! [X0] : (~relator(X0) | overlap(X0,sK30(X0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1012])])).
fof(f15395,plain,(
  spl59_1522 <=> ! [X3] : (~overlap(X3,sK30(sK41(xyz))) | externallyDependentMode(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1522])])).
fof(f15426,plain,(
  externallyDependentMode(sK41(xyz)) | (~spl59_341 | ~spl59_1012 | ~spl59_1522)),
  inference(subsumption_resolution,[],[f15398,f2278])).
fof(f2278,plain,(
  relator(sK41(xyz)) | ~spl59_341),
  inference(avatar_component_clause,[],[f2277])).
fof(f15398,plain,(
  externallyDependentMode(sK41(xyz)) | ~relator(sK41(xyz)) | (~spl59_1012 | ~spl59_1522)),
  inference(resolution,[],[f15396,f9140])).
fof(f9140,plain,(
  ( ! [X0] : (overlap(X0,sK30(X0)) | ~relator(X0)) ) | ~spl59_1012),
  inference(avatar_component_clause,[],[f9139])).
fof(f15396,plain,(
  ( ! [X3] : (~overlap(X3,sK30(sK41(xyz))) | externallyDependentMode(X3)) ) | ~spl59_1522),
  inference(avatar_component_clause,[],[f15395])).
fof(f15397,plain,(
  spl59_1522 | ~spl59_159 | ~spl59_1520),
  inference(avatar_split_clause,[],[f15392,f15383,f1283,f15395])).
fof(f1283,plain,(
  spl59_159 <=> ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_159])])).
fof(f15383,plain,(
  spl59_1520 <=> quaIndividualOf(sK30(sK41(xyz)),sK18(sK30(sK41(xyz))))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1520])])).
fof(f15392,plain,(
  ( ! [X3] : (~overlap(X3,sK30(sK41(xyz))) | externallyDependentMode(X3)) ) | (~spl59_159 | ~spl59_1520)),
  inference(resolution,[],[f15384,f1284])).
fof(f1284,plain,(
  ( ! [X2,X0,X1] : (~quaIndividualOf(X0,X1) | ~overlap(X2,X0) | externallyDependentMode(X2)) ) | ~spl59_159),
  inference(avatar_component_clause,[],[f1283])).
fof(f15384,plain,(
  quaIndividualOf(sK30(sK41(xyz)),sK18(sK30(sK41(xyz)))) | ~spl59_1520),
  inference(avatar_component_clause,[],[f15383])).
fof(f15385,plain,(
  spl59_1520 | ~spl59_94 | ~spl59_1447),
  inference(avatar_split_clause,[],[f14520,f14511,f1018,f15383])).
fof(f1018,plain,(
  spl59_94 <=> ! [X0] : (quaIndividualOf(X0,sK18(X0)) | ~quaIndividual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_94])])).
fof(f14511,plain,(
  spl59_1447 <=> quaIndividual(sK30(sK41(xyz)))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1447])])).
fof(f14520,plain,(
  quaIndividualOf(sK30(sK41(xyz)),sK18(sK30(sK41(xyz)))) | (~spl59_94 | ~spl59_1447)),
  inference(resolution,[],[f14512,f1019])).
fof(f1019,plain,(
  ( ! [X0] : (~quaIndividual(X0) | quaIndividualOf(X0,sK18(X0))) ) | ~spl59_94),
  inference(avatar_component_clause,[],[f1018])).
fof(f14512,plain,(
  quaIndividual(sK30(sK41(xyz))) | ~spl59_1447),
  inference(avatar_component_clause,[],[f14511])).
fof(f14513,plain,(
  spl59_1447 | ~spl59_99 | ~spl59_341 | ~spl59_1446),
  inference(avatar_split_clause,[],[f14509,f14500,f2277,f1038,f14511])).
fof(f1038,plain,(
  spl59_99 <=> ! [X0] : (properPartOf(sK30(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_99])])).
fof(f14500,plain,(
  spl59_1446 <=> ! [X2] : (~properPartOf(X2,sK41(xyz)) | quaIndividual(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1446])])).
fof(f14509,plain,(
  quaIndividual(sK30(sK41(xyz))) | (~spl59_99 | ~spl59_341 | ~spl59_1446)),
  inference(subsumption_resolution,[],[f14508,f2278])).
fof(f14508,plain,(
  quaIndividual(sK30(sK41(xyz))) | ~relator(sK41(xyz)) | (~spl59_99 | ~spl59_1446)),
  inference(resolution,[],[f14501,f1039])).
fof(f1039,plain,(
  ( ! [X0] : (properPartOf(sK30(X0),X0) | ~relator(X0)) ) | ~spl59_99),
  inference(avatar_component_clause,[],[f1038])).
fof(f14501,plain,(
  ( ! [X2] : (~properPartOf(X2,sK41(xyz)) | quaIndividual(X2)) ) | ~spl59_1446),
  inference(avatar_component_clause,[],[f14500])).
fof(f14506,plain,(
  ~spl59_99 | ~spl59_341 | ~spl59_1445),
  inference(avatar_contradiction_clause,[],[f14505])).
fof(f14505,plain,(
  $false | (~spl59_99 | ~spl59_341 | ~spl59_1445)),
  inference(subsumption_resolution,[],[f14504,f2278])).
fof(f14504,plain,(
  ~relator(sK41(xyz)) | (~spl59_99 | ~spl59_1445)),
  inference(resolution,[],[f14498,f1039])).
fof(f14498,plain,(
  ( ! [X3] : (~properPartOf(X3,sK41(xyz))) ) | ~spl59_1445),
  inference(avatar_component_clause,[],[f14497])).
fof(f14497,plain,(
  spl59_1445 <=> ! [X3] : ~properPartOf(X3,sK41(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1445])])).
fof(f14502,plain,(
  spl59_1445 | spl59_1446 | ~spl59_254 | ~spl59_341),
  inference(avatar_split_clause,[],[f2483,f2277,f1792,f14500,f14497])).
fof(f1792,plain,(
  spl59_254 <=> ! [X5,X4,X0] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_254])])).
fof(f2483,plain,(
  ( ! [X2,X3] : (~properPartOf(X2,sK41(xyz)) | quaIndividual(X2) | ~properPartOf(X3,sK41(xyz))) ) | (~spl59_254 | ~spl59_341)),
  inference(resolution,[],[f1793,f2278])).
fof(f1793,plain,(
  ( ! [X0,X4,X5] : (~relator(X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~properPartOf(X4,X0)) ) | ~spl59_254),
  inference(avatar_component_clause,[],[f1792])).
fof(f9141,plain,(
  spl59_1012 | ~spl59_750 | ~spl59_992),
  inference(avatar_split_clause,[],[f8877,f8874,f7046,f9139])).
fof(f7046,plain,(
  spl59_750 <=> ! [X4,X5] : (overlap(X4,X5) | ~overlap(X5,X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_750])])).
fof(f8874,plain,(
  spl59_992 <=> ! [X0] : (overlap(sK30(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_992])])).
fof(f8877,plain,(
  ( ! [X0] : (~relator(X0) | overlap(X0,sK30(X0))) ) | (~spl59_750 | ~spl59_992)),
  inference(resolution,[],[f8875,f7047])).
fof(f7047,plain,(
  ( ! [X4,X5] : (~overlap(X5,X4) | overlap(X4,X5)) ) | ~spl59_750),
  inference(avatar_component_clause,[],[f7046])).
fof(f8875,plain,(
  ( ! [X0] : (overlap(sK30(X0),X0) | ~relator(X0)) ) | ~spl59_992),
  inference(avatar_component_clause,[],[f8874])).
fof(f8876,plain,(
  spl59_992 | ~spl59_5 | ~spl59_937),
  inference(avatar_split_clause,[],[f8461,f8450,f656,f8874])).
fof(f656,plain,(
  spl59_5 <=> ! [X0] : partOf(X0,X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_5])])).
fof(f8450,plain,(
  spl59_937 <=> ! [X2,X3] : (overlap(X2,X3) | ~partOf(X2,sK30(X3)) | ~relator(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_937])])).
fof(f8461,plain,(
  ( ! [X0] : (overlap(sK30(X0),X0) | ~relator(X0)) ) | (~spl59_5 | ~spl59_937)),
  inference(resolution,[],[f8451,f657])).
fof(f657,plain,(
  ( ! [X0] : (partOf(X0,X0)) ) | ~spl59_5),
  inference(avatar_component_clause,[],[f656])).
fof(f8451,plain,(
  ( ! [X2,X3] : (~partOf(X2,sK30(X3)) | overlap(X2,X3) | ~relator(X3)) ) | ~spl59_937),
  inference(avatar_component_clause,[],[f8450])).
fof(f8452,plain,(
  spl59_937 | ~spl59_166 | ~spl59_174),
  inference(avatar_split_clause,[],[f1961,f1357,f1325,f8450])).
fof(f1325,plain,(
  spl59_166 <=> ! [X0] : (partOf(sK30(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_166])])).
fof(f1357,plain,(
  spl59_174 <=> ! [X2,X0,X1] : (~partOf(X2,X0) | overlap(X2,X1) | ~partOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_174])])).
fof(f1961,plain,(
  ( ! [X2,X3] : (overlap(X2,X3) | ~partOf(X2,sK30(X3)) | ~relator(X3)) ) | (~spl59_166 | ~spl59_174)),
  inference(resolution,[],[f1358,f1326])).
fof(f1326,plain,(
  ( ! [X0] : (partOf(sK30(X0),X0) | ~relator(X0)) ) | ~spl59_166),
  inference(avatar_component_clause,[],[f1325])).
fof(f1358,plain,(
  ( ! [X2,X0,X1] : (~partOf(X0,X1) | overlap(X2,X1) | ~partOf(X2,X0)) ) | ~spl59_174),
  inference(avatar_component_clause,[],[f1357])).
fof(f7048,plain,(
  spl59_750 | ~spl59_161 | ~spl59_717),
  inference(avatar_split_clause,[],[f6782,f6768,f1291,f7046])).
fof(f1291,plain,(
  spl59_161 <=> ! [X0,X1] : (partOf(sK52(X0,X1),X1) | ~overlap(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_161])])).
fof(f6768,plain,(
  spl59_717 <=> ! [X9,X8,X7] : (~partOf(sK52(X7,X8),X9) | overlap(X9,X7) | ~overlap(X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_717])])).
fof(f6782,plain,(
  ( ! [X4,X5] : (overlap(X4,X5) | ~overlap(X5,X4)) ) | (~spl59_161 | ~spl59_717)),
  inference(duplicate_literal_removal,[],[f6781])).
fof(f6781,plain,(
  ( ! [X4,X5] : (overlap(X4,X5) | ~overlap(X5,X4) | ~overlap(X5,X4)) ) | (~spl59_161 | ~spl59_717)),
  inference(resolution,[],[f6769,f1292])).
fof(f1292,plain,(
  ( ! [X0,X1] : (partOf(sK52(X0,X1),X1) | ~overlap(X0,X1)) ) | ~spl59_161),
  inference(avatar_component_clause,[],[f1291])).
fof(f6769,plain,(
  ( ! [X8,X9,X7] : (~partOf(sK52(X7,X8),X9) | overlap(X9,X7) | ~overlap(X7,X8)) ) | ~spl59_717),
  inference(avatar_component_clause,[],[f6768])).
fof(f6770,plain,(
  spl59_717 | ~spl59_160 | ~spl59_178),
  inference(avatar_split_clause,[],[f2007,f1373,f1287,f6768])).
fof(f1287,plain,(
  spl59_160 <=> ! [X0,X1] : (partOf(sK52(X0,X1),X0) | ~overlap(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_160])])).
fof(f1373,plain,(
  spl59_178 <=> ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_178])])).
fof(f2007,plain,(
  ( ! [X8,X9,X7] : (~partOf(sK52(X7,X8),X9) | overlap(X9,X7) | ~overlap(X7,X8)) ) | (~spl59_160 | ~spl59_178)),
  inference(resolution,[],[f1374,f1288])).
fof(f1288,plain,(
  ( ! [X0,X1] : (partOf(sK52(X0,X1),X0) | ~overlap(X0,X1)) ) | ~spl59_160),
  inference(avatar_component_clause,[],[f1287])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X1) | ~partOf(X2,X0) | overlap(X0,X1)) ) | ~spl59_178),
  inference(avatar_component_clause,[],[f1373])).
fof(f2306,plain,(
  ~spl59_344 | ~spl59_9 | ~spl59_341),
  inference(avatar_split_clause,[],[f2289,f2277,f672,f2304])).
fof(f672,plain,(
  spl59_9 <=> ! [X0] : (~intrinsicMoment(X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_9])])).
fof(f2289,plain,(
  ~intrinsicMoment(sK41(xyz)) | (~spl59_9 | ~spl59_341)),
  inference(resolution,[],[f2278,f673])).
fof(f673,plain,(
  ( ! [X0] : (~relator(X0) | ~intrinsicMoment(X0)) ) | ~spl59_9),
  inference(avatar_component_clause,[],[f672])).
fof(f2285,plain,(
  ~spl59_16 | ~spl59_84 | ~spl59_340),
  inference(avatar_contradiction_clause,[],[f2284])).
fof(f2284,plain,(
  $false | (~spl59_16 | ~spl59_84 | ~spl59_340)),
  inference(subsumption_resolution,[],[f2281,f974])).
fof(f974,plain,(
  type_(xyz) | ~spl59_84),
  inference(avatar_component_clause,[],[f973])).
fof(f973,plain,(
  spl59_84 <=> type_(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_84])])).
fof(f2281,plain,(
  ~type_(xyz) | (~spl59_16 | ~spl59_340)),
  inference(resolution,[],[f2275,f701])).
fof(f701,plain,(
  ( ! [X0] : (~individual(X0) | ~type_(X0)) ) | ~spl59_16),
  inference(avatar_component_clause,[],[f700])).
fof(f700,plain,(
  spl59_16 <=> ! [X0] : (~type_(X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_16])])).
fof(f2275,plain,(
  individual(xyz) | ~spl59_340),
  inference(avatar_component_clause,[],[f2274])).
fof(f2274,plain,(
  spl59_340 <=> individual(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_340])])).
fof(f2279,plain,(
  spl59_340 | spl59_341 | ~spl59_152 | ~spl59_315),
  inference(avatar_split_clause,[],[f2139,f2136,f1255,f2277,f2274])).
fof(f1255,plain,(
  spl59_152 <=> ! [X0] : (iof(sK41(X0),X0,sK42(X0)) | individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_152])])).
fof(f2136,plain,(
  spl59_315 <=> ! [X0,X1] : (relator(X0) | ~iof(X0,xyz,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_315])])).
fof(f2139,plain,(
  relator(sK41(xyz)) | individual(xyz) | (~spl59_152 | ~spl59_315)),
  inference(resolution,[],[f2137,f1256])).
fof(f1256,plain,(
  ( ! [X0] : (iof(sK41(X0),X0,sK42(X0)) | individual(X0)) ) | ~spl59_152),
  inference(avatar_component_clause,[],[f1255])).
fof(f2137,plain,(
  ( ! [X0,X1] : (~iof(X0,xyz,X1) | relator(X0)) ) | ~spl59_315),
  inference(avatar_component_clause,[],[f2136])).
fof(f2138,plain,(
  spl59_315 | ~spl59_2 | ~spl59_314),
  inference(avatar_split_clause,[],[f2133,f2130,f644,f2136])).
fof(f644,plain,(
  spl59_2 <=> relatorType(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_2])])).
fof(f2130,plain,(
  spl59_314 <=> ! [X2,X0,X1] : (~iof(X1,X0,X2) | relator(X1) | ~relatorType(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_314])])).
fof(f2133,plain,(
  ( ! [X0,X1] : (relator(X0) | ~iof(X0,xyz,X1)) ) | (~spl59_2 | ~spl59_314)),
  inference(resolution,[],[f2131,f645])).
fof(f645,plain,(
  relatorType(xyz) | ~spl59_2),
  inference(avatar_component_clause,[],[f644])).
fof(f2131,plain,(
  ( ! [X2,X0,X1] : (~relatorType(X0) | relator(X1) | ~iof(X1,X0,X2)) ) | ~spl59_314),
  inference(avatar_component_clause,[],[f2130])).
fof(f2132,plain,(
  spl59_314 | ~spl59_120 | ~spl59_261),
  inference(avatar_split_clause,[],[f1828,f1825,f1122,f2130])).
fof(f1122,plain,(
  spl59_120 <=> ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_120])])).
fof(f1825,plain,(
  spl59_261 <=> ! [X2,X0,X1] : (~world(X2) | ~iof(X1,X0,X2) | relator(X1) | ~relatorType(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_261])])).
fof(f1828,plain,(
  ( ! [X2,X0,X1] : (~iof(X1,X0,X2) | relator(X1) | ~relatorType(X0)) ) | (~spl59_120 | ~spl59_261)),
  inference(subsumption_resolution,[],[f1826,f1123])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2)) ) | ~spl59_120),
  inference(avatar_component_clause,[],[f1122])).
fof(f1826,plain,(
  ( ! [X2,X0,X1] : (~world(X2) | ~iof(X1,X0,X2) | relator(X1) | ~relatorType(X0)) ) | ~spl59_261),
  inference(avatar_component_clause,[],[f1825])).
fof(f1827,plain,(
  spl59_261),
  inference(avatar_split_clause,[],[f482,f1825])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~world(X2) | ~iof(X1,X0,X2) | relator(X1) | ~relatorType(X0)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0] : (relatorType(X0) <=> (! [X1,X2] : (relator(X1) | ~iof(X1,X0,X2) | ~world(X2)) & type_(X0)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0] : (relatorType(X0) <=> (! [X1,X2] : (relator(X1) | (~iof(X1,X0,X2) | ~world(X2))) & type_(X0)))),
  inference(ennf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : (relatorType(X0) <=> (! [X1,X2] : ((iof(X1,X0,X2) & world(X2)) => relator(X1)) & type_(X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X8] : (relatorType(X8) <=> (! [X6,X2] : ((iof(X6,X8,X2) & world(X2)) => relator(X6)) & type_(X8)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_relatorTypeDefinition_a44)).
fof(f1794,plain,(
  spl59_254),
  inference(avatar_split_clause,[],[f465,f1792])).
fof(f465,plain,(
  ( ! [X0,X4,X5] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | ? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0)) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | ~properPartOf(X5,X0) | ~properPartOf(X4,X0)) & ? [X7] : properPartOf(X7,X0)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | (? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0))) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | (~properPartOf(X5,X0) | ~properPartOf(X4,X0))) & ? [X7] : properPartOf(X7,X0)))),
  inference(ennf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : ((! [X3] : (foundedBy(X1,X3) <=> foundedBy(X2,X3)) & existentiallyDependsOn(X2,X1) & existentiallyDependsOn(X1,X2) & quaIndividual(X2) & properPartOf(X1,X0)) => properPartOf(X2,X0)) & ! [X4,X5] : ((properPartOf(X5,X0) & properPartOf(X4,X0)) => (! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4))) & ? [X7] : properPartOf(X7,X0)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X23] : (relator(X23) <=> (! [X24,X25] : ((! [X18] : (foundedBy(X24,X18) <=> foundedBy(X25,X18)) & existentiallyDependsOn(X25,X24) & existentiallyDependsOn(X24,X25) & quaIndividual(X25) & properPartOf(X24,X23)) => properPartOf(X25,X23)) & ! [X1,X3] : ((properPartOf(X3,X23) & properPartOf(X1,X23)) => (! [X14] : (foundedBy(X1,X14) <=> foundedBy(X3,X14)) & existentiallyDependsOn(X3,X1) & existentiallyDependsOn(X1,X3) & quaIndividual(X3) & quaIndividual(X1))) & ? [X0] : properPartOf(X0,X23)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_dRelator_a81)).
fof(f1375,plain,(
  spl59_178),
  inference(avatar_split_clause,[],[f587,f1373])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X2] : (partOf(X2,X1) & partOf(X2,X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X3] : (partOf(X3,X1) & partOf(X3,X0)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_overlappingWholes_a50)).
fof(f1359,plain,(
  spl59_174),
  inference(avatar_split_clause,[],[f568,f1357])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X0) | overlap(X2,X1) | ~partOf(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (~partOf(X0,X1) <=> ? [X2] : (~overlap(X2,X1) & partOf(X2,X0)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X0,X1] : (~partOf(X0,X1) <=> ? [X3] : (~overlap(X3,X1) & partOf(X3,X0)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_strongSupplementation_a51)).
fof(f1327,plain,(
  spl59_166 | ~spl59_99 | ~spl59_116),
  inference(avatar_split_clause,[],[f1247,f1106,f1038,f1325])).
fof(f1106,plain,(
  spl59_116 <=> ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_116])])).
fof(f1247,plain,(
  ( ! [X0] : (partOf(sK30(X0),X0) | ~relator(X0)) ) | (~spl59_99 | ~spl59_116)),
  inference(resolution,[],[f1107,f1039])).
fof(f1107,plain,(
  ( ! [X0,X1] : (~properPartOf(X0,X1) | partOf(X0,X1)) ) | ~spl59_116),
  inference(avatar_component_clause,[],[f1106])).
fof(f1293,plain,(
  spl59_161),
  inference(avatar_split_clause,[],[f586,f1291])).
fof(f586,plain,(
  ( ! [X0,X1] : (partOf(sK52(X0,X1),X1) | ~overlap(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f1289,plain,(
  spl59_160),
  inference(avatar_split_clause,[],[f585,f1287])).
fof(f585,plain,(
  ( ! [X0,X1] : (partOf(sK52(X0,X1),X0) | ~overlap(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f1285,plain,(
  spl59_159),
  inference(avatar_split_clause,[],[f580,f1283])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X2,X3) | ~foundedBy(X0,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X0,X3) => foundedBy(X2,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X3] : (overlap(X3,X0) <=> (! [X14] : (foundedBy(X0,X14) => foundedBy(X3,X14)) & inheresIn(X3,X1) & externallyDependentMode(X3))))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_dQuaIndividualOf_a75)).
fof(f1257,plain,(
  spl59_152),
  inference(avatar_split_clause,[],[f489,f1255])).
fof(f489,plain,(
  ( ! [X0] : (iof(sK41(X0),X0,sK42(X0)) | individual(X0)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0] : (individual(X0) <=> ! [X1,X2] : ~iof(X1,X0,X2))),
  inference(ennf_transformation,[],[f50])).
fof(f50,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1,X2] : iof(X1,X0,X2))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_dIndividual_a2)).
fof(f1124,plain,(
  spl59_120),
  inference(avatar_split_clause,[],[f624,f1122])).
fof(f624,plain,(
  ( ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2)) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((world(X2) & type_(X1)) | ~iof(X0,X1,X2))),
  inference(ennf_transformation,[],[f48])).
fof(f48,axiom,(
  ! [X0,X1,X2] : (iof(X0,X1,X2) => (world(X2) & type_(X1)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_dIof)).
fof(f1108,plain,(
  spl59_116),
  inference(avatar_split_clause,[],[f582,f1106])).
fof(f582,plain,(
  ( ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,axiom,(
  ! [X0,X1] : (properPartOf(X0,X1) <=> (~partOf(X1,X0) & partOf(X0,X1)))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_properPart_a52)).
fof(f1040,plain,(
  spl59_99),
  inference(avatar_split_clause,[],[f478,f1038])).
fof(f478,plain,(
  ( ! [X0] : (properPartOf(sK30(X0),X0) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f262])).
fof(f1020,plain,(
  spl59_94),
  inference(avatar_split_clause,[],[f401,f1018])).
fof(f401,plain,(
  ( ! [X0] : (quaIndividualOf(X0,sK18(X0)) | ~quaIndividual(X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,axiom,(
  ! [X0] : (quaIndividual(X0) <=> ? [X1] : quaIndividualOf(X0,X1))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_dQuaIndividual_a76)).
fof(f975,plain,(
  spl59_84 | ~spl59_2 | ~spl59_49),
  inference(avatar_split_clause,[],[f971,f832,f644,f973])).
fof(f832,plain,(
  spl59_49 <=> ! [X0] : (type_(X0) | ~relatorType(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_49])])).
fof(f971,plain,(
  type_(xyz) | (~spl59_2 | ~spl59_49)),
  inference(resolution,[],[f833,f645])).
fof(f833,plain,(
  ( ! [X0] : (~relatorType(X0) | type_(X0)) ) | ~spl59_49),
  inference(avatar_component_clause,[],[f832])).
fof(f846,plain,(
  spl59_52),
  inference(avatar_split_clause,[],[f511,f844])).
fof(f511,plain,(
  ( ! [X0] : (intrinsicMoment(X0) | ~mode(X0)) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  ! [X0] : ((mode(X0) | quality(X0)) <=> intrinsicMoment(X0))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_intrinsicMoment_taxonomy)).
fof(f834,plain,(
  spl59_49),
  inference(avatar_split_clause,[],[f483,f832])).
fof(f483,plain,(
  ( ! [X0] : (type_(X0) | ~relatorType(X0)) )),
  inference(cnf_transformation,[],[f264])).
fof(f746,plain,(
  spl59_27),
  inference(avatar_split_clause,[],[f345,f744])).
fof(f345,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (mode(X0) | ~externallyDependentMode(X0))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0] : (externallyDependentMode(X0) => mode(X0))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_mode_taxonomy_externallyDependentMode)).
fof(f702,plain,(
  spl59_16),
  inference(avatar_split_clause,[],[f325,f700])).
fof(f325,plain,(
  ( ! [X0] : (~type_(X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0] : (~individual(X0) | ~type_(X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ~? [X0] : (individual(X0) & type_(X0))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_thing_partition)).
fof(f674,plain,(
  spl59_9),
  inference(avatar_split_clause,[],[f316,f672])).
fof(f316,plain,(
  ( ! [X0] : (~intrinsicMoment(X0) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0] : (~relator(X0) | ~intrinsicMoment(X0))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ~? [X0] : (relator(X0) & intrinsicMoment(X0))),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_moment_partition)).
fof(f658,plain,(
  spl59_5),
  inference(avatar_split_clause,[],[f342,f656])).
fof(f342,plain,(
  ( ! [X0] : (partOf(X0,X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  ! [X0] : partOf(X0,X0)),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',ax_reflexiveParthood)).
fof(f646,plain,(
  spl59_2),
  inference(avatar_split_clause,[],[f308,f644])).
fof(f308,plain,(
  relatorType(xyz)),
  inference(cnf_transformation,[],[f132])).
fof(f132,axiom,(
  relatorType(xyz)),
  file('outputs/ufo/subsumptions/type_relatorType_subsumed_to_kind.in',axiom1)).
% SZS output end Proof for type_relatorType_subsumed_to_kind
% ------------------------------
% Version: Vampire 4.7 (commit af7d05f97 on 2022-09-30 11:01:41 +0200)
% Termination reason: Refutation

% Memory used [KB]: 19829
% Time elapsed: 0.379 s
% ------------------------------
% ------------------------------
% Success in time 0.389 s
