% Running in auto input_syntax mode. Trying TPTP
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on type_relator_subsumed_to_antiRigidSortal for (31ds)
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for type_relator_subsumed_to_antiRigidSortal
% SZS output start Proof for type_relator_subsumed_to_antiRigidSortal
fof(f6146,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f646,f658,f686,f746,f846,f1024,f1056,f1116,f1376,f1477,f1600,f2014,f2118,f3020,f4280,f4287,f4974,f6093,f6102,f6121,f6128,f6137,f6145])).
fof(f6145,plain,(
  spl59_351 | ~spl59_52 | ~spl59_779),
  inference(avatar_split_clause,[],[f6143,f6135,f844,f2993])).
fof(f2993,plain,(
  spl59_351 <=> intrinsicMoment(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_351])])).
fof(f844,plain,(
  spl59_52 <=> ! [X0] : (intrinsicMoment(X0) | ~mode(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_52])])).
fof(f6135,plain,(
  spl59_779 <=> mode(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_779])])).
fof(f6143,plain,(
  intrinsicMoment(xyz) | (~spl59_52 | ~spl59_779)),
  inference(resolution,[],[f6136,f845])).
fof(f845,plain,(
  ( ! [X0] : (~mode(X0) | intrinsicMoment(X0)) ) | ~spl59_52),
  inference(avatar_component_clause,[],[f844])).
fof(f6136,plain,(
  mode(xyz) | ~spl59_779),
  inference(avatar_component_clause,[],[f6135])).
fof(f6137,plain,(
  spl59_779 | ~spl59_27 | ~spl59_778),
  inference(avatar_split_clause,[],[f6133,f6126,f744,f6135])).
fof(f744,plain,(
  spl59_27 <=> ! [X0] : (~externallyDependentMode(X0) | mode(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_27])])).
fof(f6126,plain,(
  spl59_778 <=> externallyDependentMode(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_778])])).
fof(f6133,plain,(
  mode(xyz) | (~spl59_27 | ~spl59_778)),
  inference(resolution,[],[f6127,f745])).
fof(f745,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) ) | ~spl59_27),
  inference(avatar_component_clause,[],[f744])).
fof(f6127,plain,(
  externallyDependentMode(xyz) | ~spl59_778),
  inference(avatar_component_clause,[],[f6126])).
fof(f6128,plain,(
  spl59_778 | ~spl59_2 | ~spl59_192 | ~spl59_777),
  inference(avatar_split_clause,[],[f6124,f6119,f1475,f644,f6126])).
fof(f644,plain,(
  spl59_2 <=> relator(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_2])])).
fof(f1475,plain,(
  spl59_192 <=> ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_192])])).
fof(f6119,plain,(
  spl59_777 <=> ! [X0] : (externallyDependentMode(X0) | ~partOf(sK39(xyz),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_777])])).
fof(f6124,plain,(
  externallyDependentMode(xyz) | (~spl59_2 | ~spl59_192 | ~spl59_777)),
  inference(subsumption_resolution,[],[f6122,f645])).
fof(f645,plain,(
  relator(xyz) | ~spl59_2),
  inference(avatar_component_clause,[],[f644])).
fof(f6122,plain,(
  externallyDependentMode(xyz) | ~relator(xyz) | (~spl59_192 | ~spl59_777)),
  inference(resolution,[],[f6120,f1476])).
fof(f1476,plain,(
  ( ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0)) ) | ~spl59_192),
  inference(avatar_component_clause,[],[f1475])).
fof(f6120,plain,(
  ( ! [X0] : (~partOf(sK39(xyz),X0) | externallyDependentMode(X0)) ) | ~spl59_777),
  inference(avatar_component_clause,[],[f6119])).
fof(f6121,plain,(
  spl59_777 | ~spl59_320 | ~spl59_775),
  inference(avatar_split_clause,[],[f6108,f6100,f2116,f6119])).
fof(f2116,plain,(
  spl59_320 <=> ! [X0,X1] : (~partOf(X0,X1) | overlap(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_320])])).
fof(f6100,plain,(
  spl59_775 <=> ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_775])])).
fof(f6108,plain,(
  ( ! [X0] : (externallyDependentMode(X0) | ~partOf(sK39(xyz),X0)) ) | (~spl59_320 | ~spl59_775)),
  inference(resolution,[],[f6101,f2117])).
fof(f2117,plain,(
  ( ! [X0,X1] : (overlap(X1,X0) | ~partOf(X0,X1)) ) | ~spl59_320),
  inference(avatar_component_clause,[],[f2116])).
fof(f6101,plain,(
  ( ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2)) ) | ~spl59_775),
  inference(avatar_component_clause,[],[f6100])).
fof(f6102,plain,(
  spl59_775 | ~spl59_2 | ~spl59_774),
  inference(avatar_split_clause,[],[f6095,f6091,f644,f6100])).
fof(f6091,plain,(
  spl59_774 <=> ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_774])])).
fof(f6095,plain,(
  ( ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2)) ) | (~spl59_2 | ~spl59_774)),
  inference(resolution,[],[f6092,f645])).
fof(f6092,plain,(
  ( ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6)) ) | ~spl59_774),
  inference(avatar_component_clause,[],[f6091])).
fof(f6093,plain,(
  spl59_774 | ~spl59_175 | ~spl59_621),
  inference(avatar_split_clause,[],[f4977,f4972,f1374,f6091])).
fof(f1374,plain,(
  spl59_175 <=> ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_175])])).
fof(f4972,plain,(
  spl59_621 <=> ! [X0] : (~relator(X0) | quaIndividualOf(sK39(X0),sK20(sK39(X0))))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_621])])).
fof(f4977,plain,(
  ( ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6)) ) | (~spl59_175 | ~spl59_621)),
  inference(resolution,[],[f4973,f1375])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (~quaIndividualOf(X0,X1) | ~overlap(X2,X0) | externallyDependentMode(X2)) ) | ~spl59_175),
  inference(avatar_component_clause,[],[f1374])).
fof(f4973,plain,(
  ( ! [X0] : (quaIndividualOf(sK39(X0),sK20(sK39(X0))) | ~relator(X0)) ) | ~spl59_621),
  inference(avatar_component_clause,[],[f4972])).
fof(f4974,plain,(
  spl59_621 | ~spl59_95 | ~spl59_526),
  inference(avatar_split_clause,[],[f4289,f4285,f1022,f4972])).
fof(f1022,plain,(
  spl59_95 <=> ! [X0] : (quaIndividualOf(X0,sK20(X0)) | ~quaIndividual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_95])])).
fof(f4285,plain,(
  spl59_526 <=> ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_526])])).
fof(f4289,plain,(
  ( ! [X0] : (~relator(X0) | quaIndividualOf(sK39(X0),sK20(sK39(X0)))) ) | (~spl59_95 | ~spl59_526)),
  inference(resolution,[],[f4286,f1023])).
fof(f1023,plain,(
  ( ! [X0] : (~quaIndividual(X0) | quaIndividualOf(X0,sK20(X0))) ) | ~spl59_95),
  inference(avatar_component_clause,[],[f1022])).
fof(f4286,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2)) ) | ~spl59_526),
  inference(avatar_component_clause,[],[f4285])).
fof(f4287,plain,(
  spl59_526 | ~spl59_103 | ~spl59_525),
  inference(avatar_split_clause,[],[f4283,f4278,f1054,f4285])).
fof(f1054,plain,(
  spl59_103 <=> ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_103])])).
fof(f4278,plain,(
  spl59_525 <=> ! [X4,X3] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_525])])).
fof(f4283,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2)) ) | (~spl59_103 | ~spl59_525)),
  inference(duplicate_literal_removal,[],[f4282])).
fof(f4282,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2) | ~relator(X2)) ) | (~spl59_103 | ~spl59_525)),
  inference(resolution,[],[f4279,f1055])).
fof(f1055,plain,(
  ( ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0)) ) | ~spl59_103),
  inference(avatar_component_clause,[],[f1054])).
fof(f4279,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4)) ) | ~spl59_525),
  inference(avatar_component_clause,[],[f4278])).
fof(f4280,plain,(
  spl59_525 | ~spl59_103 | ~spl59_301),
  inference(avatar_split_clause,[],[f3427,f2012,f1054,f4278])).
fof(f2012,plain,(
  spl59_301 <=> ! [X5,X4,X0] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_301])])).
fof(f3427,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4)) ) | (~spl59_103 | ~spl59_301)),
  inference(duplicate_literal_removal,[],[f3426])).
fof(f3426,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4) | ~relator(X4)) ) | (~spl59_103 | ~spl59_301)),
  inference(resolution,[],[f2013,f1055])).
fof(f2013,plain,(
  ( ! [X0,X4,X5] : (~properPartOf(X5,X0) | ~properPartOf(X4,X0) | quaIndividual(X5) | ~relator(X0)) ) | ~spl59_301),
  inference(avatar_component_clause,[],[f2012])).
fof(f3020,plain,(
  ~spl59_2 | ~spl59_12 | ~spl59_351),
  inference(avatar_contradiction_clause,[],[f3019])).
fof(f3019,plain,(
  $false | (~spl59_2 | ~spl59_12 | ~spl59_351)),
  inference(subsumption_resolution,[],[f3018,f645])).
fof(f3018,plain,(
  ~relator(xyz) | (~spl59_12 | ~spl59_351)),
  inference(resolution,[],[f2994,f685])).
fof(f685,plain,(
  ( ! [X0] : (~intrinsicMoment(X0) | ~relator(X0)) ) | ~spl59_12),
  inference(avatar_component_clause,[],[f684])).
fof(f684,plain,(
  spl59_12 <=> ! [X0] : (~intrinsicMoment(X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_12])])).
fof(f2994,plain,(
  intrinsicMoment(xyz) | ~spl59_351),
  inference(avatar_component_clause,[],[f2993])).
fof(f2118,plain,(
  spl59_320 | ~spl59_5 | ~spl59_217),
  inference(avatar_split_clause,[],[f2110,f1598,f656,f2116])).
fof(f656,plain,(
  spl59_5 <=> ! [X0] : partOf(X0,X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_5])])).
fof(f1598,plain,(
  spl59_217 <=> ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_217])])).
fof(f2110,plain,(
  ( ! [X0,X1] : (~partOf(X0,X1) | overlap(X1,X0)) ) | (~spl59_5 | ~spl59_217)),
  inference(resolution,[],[f1599,f657])).
fof(f657,plain,(
  ( ! [X0] : (partOf(X0,X0)) ) | ~spl59_5),
  inference(avatar_component_clause,[],[f656])).
fof(f1599,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X1) | ~partOf(X2,X0) | overlap(X0,X1)) ) | ~spl59_217),
  inference(avatar_component_clause,[],[f1598])).
fof(f2014,plain,(
  spl59_301),
  inference(avatar_split_clause,[],[f485,f2012])).
fof(f485,plain,(
  ( ! [X0,X4,X5] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | ? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0)) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | ~properPartOf(X5,X0) | ~properPartOf(X4,X0)) & ? [X7] : properPartOf(X7,X0)))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | (? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0))) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | (~properPartOf(X5,X0) | ~properPartOf(X4,X0))) & ? [X7] : properPartOf(X7,X0)))),
  inference(ennf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : ((! [X3] : (foundedBy(X1,X3) <=> foundedBy(X2,X3)) & existentiallyDependsOn(X2,X1) & existentiallyDependsOn(X1,X2) & quaIndividual(X2) & properPartOf(X1,X0)) => properPartOf(X2,X0)) & ! [X4,X5] : ((properPartOf(X5,X0) & properPartOf(X4,X0)) => (! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4))) & ? [X7] : properPartOf(X7,X0)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X23] : (relator(X23) <=> (! [X24,X25] : ((! [X18] : (foundedBy(X24,X18) <=> foundedBy(X25,X18)) & existentiallyDependsOn(X25,X24) & existentiallyDependsOn(X24,X25) & quaIndividual(X25) & properPartOf(X24,X23)) => properPartOf(X25,X23)) & ! [X1,X3] : ((properPartOf(X3,X23) & properPartOf(X1,X23)) => (! [X14] : (foundedBy(X1,X14) <=> foundedBy(X3,X14)) & existentiallyDependsOn(X3,X1) & existentiallyDependsOn(X1,X3) & quaIndividual(X3) & quaIndividual(X1))) & ? [X0] : properPartOf(X0,X23)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_dRelator_a81)).
fof(f1600,plain,(
  spl59_217),
  inference(avatar_split_clause,[],[f587,f1598])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X2] : (partOf(X2,X1) & partOf(X2,X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X3] : (partOf(X3,X1) & partOf(X3,X0)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_overlappingWholes_a50)).
fof(f1477,plain,(
  spl59_192 | ~spl59_103 | ~spl59_118),
  inference(avatar_split_clause,[],[f1258,f1114,f1054,f1475])).
fof(f1114,plain,(
  spl59_118 <=> ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_118])])).
fof(f1258,plain,(
  ( ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0)) ) | (~spl59_103 | ~spl59_118)),
  inference(resolution,[],[f1115,f1055])).
fof(f1115,plain,(
  ( ! [X0,X1] : (~properPartOf(X0,X1) | partOf(X0,X1)) ) | ~spl59_118),
  inference(avatar_component_clause,[],[f1114])).
fof(f1376,plain,(
  spl59_175),
  inference(avatar_split_clause,[],[f580,f1374])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X2,X3) | ~foundedBy(X0,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X0,X3) => foundedBy(X2,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X3] : (overlap(X3,X0) <=> (! [X14] : (foundedBy(X0,X14) => foundedBy(X3,X14)) & inheresIn(X3,X1) & externallyDependentMode(X3))))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_dQuaIndividualOf_a75)).
fof(f1116,plain,(
  spl59_118),
  inference(avatar_split_clause,[],[f582,f1114])).
fof(f582,plain,(
  ( ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,axiom,(
  ! [X0,X1] : (properPartOf(X0,X1) <=> (~partOf(X1,X0) & partOf(X0,X1)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_properPart_a52)).
fof(f1056,plain,(
  spl59_103),
  inference(avatar_split_clause,[],[f498,f1054])).
fof(f498,plain,(
  ( ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f1024,plain,(
  spl59_95),
  inference(avatar_split_clause,[],[f403,f1022])).
fof(f403,plain,(
  ( ! [X0] : (quaIndividualOf(X0,sK20(X0)) | ~quaIndividual(X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,axiom,(
  ! [X0] : (quaIndividual(X0) <=> ? [X1] : quaIndividualOf(X0,X1))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_dQuaIndividual_a76)).
fof(f846,plain,(
  spl59_52),
  inference(avatar_split_clause,[],[f511,f844])).
fof(f511,plain,(
  ( ! [X0] : (intrinsicMoment(X0) | ~mode(X0)) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  ! [X0] : ((mode(X0) | quality(X0)) <=> intrinsicMoment(X0))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_intrinsicMoment_taxonomy)).
fof(f746,plain,(
  spl59_27),
  inference(avatar_split_clause,[],[f345,f744])).
fof(f345,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (mode(X0) | ~externallyDependentMode(X0))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0] : (externallyDependentMode(X0) => mode(X0))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_mode_taxonomy_externallyDependentMode)).
fof(f686,plain,(
  spl59_12),
  inference(avatar_split_clause,[],[f319,f684])).
fof(f319,plain,(
  ( ! [X0] : (~intrinsicMoment(X0) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : (~relator(X0) | ~intrinsicMoment(X0))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ~? [X0] : (relator(X0) & intrinsicMoment(X0))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_moment_partition)).
fof(f658,plain,(
  spl59_5),
  inference(avatar_split_clause,[],[f342,f656])).
fof(f342,plain,(
  ( ! [X0] : (partOf(X0,X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  ! [X0] : partOf(X0,X0)),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',ax_reflexiveParthood)).
fof(f646,plain,(
  spl59_2),
  inference(avatar_split_clause,[],[f308,f644])).
fof(f308,plain,(
  relator(xyz)),
  inference(cnf_transformation,[],[f132])).
fof(f132,axiom,(
  relator(xyz)),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_antiRigidSortal.in',axiom1)).
% SZS output end Proof for type_relator_subsumed_to_antiRigidSortal
% ------------------------------
% Version: Vampire 4.7 (commit af7d05f97 on 2022-09-30 11:01:41 +0200)
% Termination reason: Refutation

% Memory used [KB]: 10362
% Time elapsed: 0.101 s
% ------------------------------
% ------------------------------
% Success in time 0.089 s
