% Running in auto input_syntax mode. Trying TPTP
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on type_relator_subsumed_to_mode for (31ds)
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for type_relator_subsumed_to_mode
% SZS output start Proof for type_relator_subsumed_to_mode
fof(f6129,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f642,f646,f658,f746,f1028,f1056,f1116,f1376,f1477,f1600,f2014,f2118,f4281,f4288,f4969,f6078,f6091,f6110,f6117,f6128])).
fof(f6128,plain,(
  spl59_1 | ~spl59_27 | ~spl59_775),
  inference(avatar_contradiction_clause,[],[f6127])).
fof(f6127,plain,(
  $false | (spl59_1 | ~spl59_27 | ~spl59_775)),
  inference(subsumption_resolution,[],[f6126,f641])).
fof(f641,plain,(
  ~mode(xyz) | spl59_1),
  inference(avatar_component_clause,[],[f640])).
fof(f640,plain,(
  spl59_1 <=> mode(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1])])).
fof(f6126,plain,(
  mode(xyz) | (~spl59_27 | ~spl59_775)),
  inference(resolution,[],[f6116,f745])).
fof(f745,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) ) | ~spl59_27),
  inference(avatar_component_clause,[],[f744])).
fof(f744,plain,(
  spl59_27 <=> ! [X0] : (~externallyDependentMode(X0) | mode(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_27])])).
fof(f6116,plain,(
  externallyDependentMode(xyz) | ~spl59_775),
  inference(avatar_component_clause,[],[f6115])).
fof(f6115,plain,(
  spl59_775 <=> externallyDependentMode(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_775])])).
fof(f6117,plain,(
  spl59_775 | ~spl59_2 | ~spl59_192 | ~spl59_774),
  inference(avatar_split_clause,[],[f6113,f6108,f1475,f644,f6115])).
fof(f644,plain,(
  spl59_2 <=> relator(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_2])])).
fof(f1475,plain,(
  spl59_192 <=> ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_192])])).
fof(f6108,plain,(
  spl59_774 <=> ! [X0] : (externallyDependentMode(X0) | ~partOf(sK39(xyz),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_774])])).
fof(f6113,plain,(
  externallyDependentMode(xyz) | (~spl59_2 | ~spl59_192 | ~spl59_774)),
  inference(subsumption_resolution,[],[f6111,f645])).
fof(f645,plain,(
  relator(xyz) | ~spl59_2),
  inference(avatar_component_clause,[],[f644])).
fof(f6111,plain,(
  externallyDependentMode(xyz) | ~relator(xyz) | (~spl59_192 | ~spl59_774)),
  inference(resolution,[],[f6109,f1476])).
fof(f1476,plain,(
  ( ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0)) ) | ~spl59_192),
  inference(avatar_component_clause,[],[f1475])).
fof(f6109,plain,(
  ( ! [X0] : (~partOf(sK39(xyz),X0) | externallyDependentMode(X0)) ) | ~spl59_774),
  inference(avatar_component_clause,[],[f6108])).
fof(f6110,plain,(
  spl59_774 | ~spl59_320 | ~spl59_773),
  inference(avatar_split_clause,[],[f6093,f6089,f2116,f6108])).
fof(f2116,plain,(
  spl59_320 <=> ! [X0,X1] : (~partOf(X0,X1) | overlap(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_320])])).
fof(f6089,plain,(
  spl59_773 <=> ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_773])])).
fof(f6093,plain,(
  ( ! [X0] : (externallyDependentMode(X0) | ~partOf(sK39(xyz),X0)) ) | (~spl59_320 | ~spl59_773)),
  inference(resolution,[],[f6090,f2117])).
fof(f2117,plain,(
  ( ! [X0,X1] : (overlap(X1,X0) | ~partOf(X0,X1)) ) | ~spl59_320),
  inference(avatar_component_clause,[],[f2116])).
fof(f6090,plain,(
  ( ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2)) ) | ~spl59_773),
  inference(avatar_component_clause,[],[f6089])).
fof(f6091,plain,(
  spl59_773 | ~spl59_2 | ~spl59_771),
  inference(avatar_split_clause,[],[f6084,f6076,f644,f6089])).
fof(f6076,plain,(
  spl59_771 <=> ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_771])])).
fof(f6084,plain,(
  ( ! [X2] : (~overlap(X2,sK39(xyz)) | externallyDependentMode(X2)) ) | (~spl59_2 | ~spl59_771)),
  inference(resolution,[],[f6077,f645])).
fof(f6077,plain,(
  ( ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6)) ) | ~spl59_771),
  inference(avatar_component_clause,[],[f6076])).
fof(f6078,plain,(
  spl59_771 | ~spl59_175 | ~spl59_620),
  inference(avatar_split_clause,[],[f4972,f4967,f1374,f6076])).
fof(f1374,plain,(
  spl59_175 <=> ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_175])])).
fof(f4967,plain,(
  spl59_620 <=> ! [X0] : (~relator(X0) | quaIndividualOf(sK39(X0),sK20(sK39(X0))))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_620])])).
fof(f4972,plain,(
  ( ! [X6,X5] : (~relator(X5) | ~overlap(X6,sK39(X5)) | externallyDependentMode(X6)) ) | (~spl59_175 | ~spl59_620)),
  inference(resolution,[],[f4968,f1375])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (~quaIndividualOf(X0,X1) | ~overlap(X2,X0) | externallyDependentMode(X2)) ) | ~spl59_175),
  inference(avatar_component_clause,[],[f1374])).
fof(f4968,plain,(
  ( ! [X0] : (quaIndividualOf(sK39(X0),sK20(sK39(X0))) | ~relator(X0)) ) | ~spl59_620),
  inference(avatar_component_clause,[],[f4967])).
fof(f4969,plain,(
  spl59_620 | ~spl59_96 | ~spl59_526),
  inference(avatar_split_clause,[],[f4290,f4286,f1026,f4967])).
fof(f1026,plain,(
  spl59_96 <=> ! [X0] : (quaIndividualOf(X0,sK20(X0)) | ~quaIndividual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_96])])).
fof(f4286,plain,(
  spl59_526 <=> ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_526])])).
fof(f4290,plain,(
  ( ! [X0] : (~relator(X0) | quaIndividualOf(sK39(X0),sK20(sK39(X0)))) ) | (~spl59_96 | ~spl59_526)),
  inference(resolution,[],[f4287,f1027])).
fof(f1027,plain,(
  ( ! [X0] : (~quaIndividual(X0) | quaIndividualOf(X0,sK20(X0))) ) | ~spl59_96),
  inference(avatar_component_clause,[],[f1026])).
fof(f4287,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2)) ) | ~spl59_526),
  inference(avatar_component_clause,[],[f4286])).
fof(f4288,plain,(
  spl59_526 | ~spl59_103 | ~spl59_525),
  inference(avatar_split_clause,[],[f4284,f4279,f1054,f4286])).
fof(f1054,plain,(
  spl59_103 <=> ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_103])])).
fof(f4279,plain,(
  spl59_525 <=> ! [X4,X3] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_525])])).
fof(f4284,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2)) ) | (~spl59_103 | ~spl59_525)),
  inference(duplicate_literal_removal,[],[f4283])).
fof(f4283,plain,(
  ( ! [X2] : (quaIndividual(sK39(X2)) | ~relator(X2) | ~relator(X2)) ) | (~spl59_103 | ~spl59_525)),
  inference(resolution,[],[f4280,f1055])).
fof(f1055,plain,(
  ( ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0)) ) | ~spl59_103),
  inference(avatar_component_clause,[],[f1054])).
fof(f4280,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4)) ) | ~spl59_525),
  inference(avatar_component_clause,[],[f4279])).
fof(f4281,plain,(
  spl59_525 | ~spl59_103 | ~spl59_301),
  inference(avatar_split_clause,[],[f3428,f2012,f1054,f4279])).
fof(f2012,plain,(
  spl59_301 <=> ! [X5,X4,X0] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_301])])).
fof(f3428,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4)) ) | (~spl59_103 | ~spl59_301)),
  inference(duplicate_literal_removal,[],[f3427])).
fof(f3427,plain,(
  ( ! [X3,X4] : (~properPartOf(X3,X4) | quaIndividual(sK39(X4)) | ~relator(X4) | ~relator(X4)) ) | (~spl59_103 | ~spl59_301)),
  inference(resolution,[],[f2013,f1055])).
fof(f2013,plain,(
  ( ! [X0,X4,X5] : (~properPartOf(X5,X0) | ~properPartOf(X4,X0) | quaIndividual(X5) | ~relator(X0)) ) | ~spl59_301),
  inference(avatar_component_clause,[],[f2012])).
fof(f2118,plain,(
  spl59_320 | ~spl59_5 | ~spl59_217),
  inference(avatar_split_clause,[],[f2110,f1598,f656,f2116])).
fof(f656,plain,(
  spl59_5 <=> ! [X0] : partOf(X0,X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_5])])).
fof(f1598,plain,(
  spl59_217 <=> ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_217])])).
fof(f2110,plain,(
  ( ! [X0,X1] : (~partOf(X0,X1) | overlap(X1,X0)) ) | (~spl59_5 | ~spl59_217)),
  inference(resolution,[],[f1599,f657])).
fof(f657,plain,(
  ( ! [X0] : (partOf(X0,X0)) ) | ~spl59_5),
  inference(avatar_component_clause,[],[f656])).
fof(f1599,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X1) | ~partOf(X2,X0) | overlap(X0,X1)) ) | ~spl59_217),
  inference(avatar_component_clause,[],[f1598])).
fof(f2014,plain,(
  spl59_301),
  inference(avatar_split_clause,[],[f485,f2012])).
fof(f485,plain,(
  ( ! [X0,X4,X5] : (~properPartOf(X4,X0) | ~properPartOf(X5,X0) | quaIndividual(X5) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | ? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0)) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | ~properPartOf(X5,X0) | ~properPartOf(X4,X0)) & ? [X7] : properPartOf(X7,X0)))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : (properPartOf(X2,X0) | (? [X3] : (foundedBy(X1,X3) <~> foundedBy(X2,X3)) | ~existentiallyDependsOn(X2,X1) | ~existentiallyDependsOn(X1,X2) | ~quaIndividual(X2) | ~properPartOf(X1,X0))) & ! [X4,X5] : ((! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4)) | (~properPartOf(X5,X0) | ~properPartOf(X4,X0))) & ? [X7] : properPartOf(X7,X0)))),
  inference(ennf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (relator(X0) <=> (! [X1,X2] : ((! [X3] : (foundedBy(X1,X3) <=> foundedBy(X2,X3)) & existentiallyDependsOn(X2,X1) & existentiallyDependsOn(X1,X2) & quaIndividual(X2) & properPartOf(X1,X0)) => properPartOf(X2,X0)) & ! [X4,X5] : ((properPartOf(X5,X0) & properPartOf(X4,X0)) => (! [X6] : (foundedBy(X4,X6) <=> foundedBy(X5,X6)) & existentiallyDependsOn(X5,X4) & existentiallyDependsOn(X4,X5) & quaIndividual(X5) & quaIndividual(X4))) & ? [X7] : properPartOf(X7,X0)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X23] : (relator(X23) <=> (! [X24,X25] : ((! [X18] : (foundedBy(X24,X18) <=> foundedBy(X25,X18)) & existentiallyDependsOn(X25,X24) & existentiallyDependsOn(X24,X25) & quaIndividual(X25) & properPartOf(X24,X23)) => properPartOf(X25,X23)) & ! [X1,X3] : ((properPartOf(X3,X23) & properPartOf(X1,X23)) => (! [X14] : (foundedBy(X1,X14) <=> foundedBy(X3,X14)) & existentiallyDependsOn(X3,X1) & existentiallyDependsOn(X1,X3) & quaIndividual(X3) & quaIndividual(X1))) & ? [X0] : properPartOf(X0,X23)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_dRelator_a81)).
fof(f1600,plain,(
  spl59_217),
  inference(avatar_split_clause,[],[f587,f1598])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (~partOf(X2,X0) | ~partOf(X2,X1) | overlap(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X2] : (partOf(X2,X1) & partOf(X2,X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X0,X1] : (overlap(X0,X1) <=> ? [X3] : (partOf(X3,X1) & partOf(X3,X0)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_overlappingWholes_a50)).
fof(f1477,plain,(
  spl59_192 | ~spl59_103 | ~spl59_118),
  inference(avatar_split_clause,[],[f1258,f1114,f1054,f1475])).
fof(f1114,plain,(
  spl59_118 <=> ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_118])])).
fof(f1258,plain,(
  ( ! [X0] : (partOf(sK39(X0),X0) | ~relator(X0)) ) | (~spl59_103 | ~spl59_118)),
  inference(resolution,[],[f1115,f1055])).
fof(f1115,plain,(
  ( ! [X0,X1] : (~properPartOf(X0,X1) | partOf(X0,X1)) ) | ~spl59_118),
  inference(avatar_component_clause,[],[f1114])).
fof(f1376,plain,(
  spl59_175),
  inference(avatar_split_clause,[],[f580,f1374])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (externallyDependentMode(X2) | ~overlap(X2,X0) | ~quaIndividualOf(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X2,X3) | ~foundedBy(X0,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X2] : (overlap(X2,X0) <=> (! [X3] : (foundedBy(X0,X3) => foundedBy(X2,X3)) & inheresIn(X2,X1) & externallyDependentMode(X2))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X0,X1] : (quaIndividualOf(X0,X1) <=> ! [X3] : (overlap(X3,X0) <=> (! [X14] : (foundedBy(X0,X14) => foundedBy(X3,X14)) & inheresIn(X3,X1) & externallyDependentMode(X3))))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_dQuaIndividualOf_a75)).
fof(f1116,plain,(
  spl59_118),
  inference(avatar_split_clause,[],[f582,f1114])).
fof(f582,plain,(
  ( ! [X0,X1] : (partOf(X0,X1) | ~properPartOf(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,axiom,(
  ! [X0,X1] : (properPartOf(X0,X1) <=> (~partOf(X1,X0) & partOf(X0,X1)))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_properPart_a52)).
fof(f1056,plain,(
  spl59_103),
  inference(avatar_split_clause,[],[f498,f1054])).
fof(f498,plain,(
  ( ! [X0] : (properPartOf(sK39(X0),X0) | ~relator(X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f1028,plain,(
  spl59_96),
  inference(avatar_split_clause,[],[f406,f1026])).
fof(f406,plain,(
  ( ! [X0] : (quaIndividualOf(X0,sK20(X0)) | ~quaIndividual(X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,axiom,(
  ! [X0] : (quaIndividual(X0) <=> ? [X1] : quaIndividualOf(X0,X1))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_dQuaIndividual_a76)).
fof(f746,plain,(
  spl59_27),
  inference(avatar_split_clause,[],[f345,f744])).
fof(f345,plain,(
  ( ! [X0] : (~externallyDependentMode(X0) | mode(X0)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (mode(X0) | ~externallyDependentMode(X0))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0] : (externallyDependentMode(X0) => mode(X0))),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_mode_taxonomy_externallyDependentMode)).
fof(f658,plain,(
  spl59_5),
  inference(avatar_split_clause,[],[f342,f656])).
fof(f342,plain,(
  ( ! [X0] : (partOf(X0,X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  ! [X0] : partOf(X0,X0)),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',ax_reflexiveParthood)).
fof(f646,plain,(
  spl59_2),
  inference(avatar_split_clause,[],[f308,f644])).
fof(f308,plain,(
  relator(xyz)),
  inference(cnf_transformation,[],[f132])).
fof(f132,axiom,(
  relator(xyz)),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',axiom1)).
fof(f642,plain,(
  ~spl59_1),
  inference(avatar_split_clause,[],[f307,f640])).
fof(f307,plain,(
  ~mode(xyz)),
  inference(cnf_transformation,[],[f133])).
fof(f133,axiom,(
  ~mode(xyz)),
  file('outputs/ufo/subsumptions/type_relator_subsumed_to_mode.in',axiom2)).
% SZS output end Proof for type_relator_subsumed_to_mode
% ------------------------------
% Version: Vampire 4.7 (commit af7d05f97 on 2022-09-30 11:01:41 +0200)
% Termination reason: Refutation

% Memory used [KB]: 10234
% Time elapsed: 0.080 s
% ------------------------------
% ------------------------------
% Success in time 0.087 s
