% Running in auto input_syntax mode. Trying TPTP
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on ci_ufo_kind_specializes_@_2 for (31ds)
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for ci_ufo_kind_specializes_@_2
% SZS output start Proof for ci_ufo_kind_specializes_@_2
fof(f36658,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f648,f652,f656,f764,f796,f816,f932,f952,f968,f992,f1012,f1035,f1048,f1064,f1100,f1172,f1301,f1378,f2845,f2914,f3185,f5342,f5780,f13702,f26762,f27747,f27812,f27825,f27928,f28128,f36583,f36654])).
fof(f36654,plain,(
  ~spl59_1 | ~spl59_2 | spl59_3 | ~spl59_98 | ~spl59_749 | ~spl59_1701),
  inference(avatar_contradiction_clause,[],[f36653])).
fof(f36653,plain,(
  $false | (~spl59_1 | ~spl59_2 | spl59_3 | ~spl59_98 | ~spl59_749 | ~spl59_1701)),
  inference(subsumption_resolution,[],[f36652,f1047])).
fof(f1047,plain,(
  type_(xyz) | ~spl59_98),
  inference(avatar_component_clause,[],[f1046])).
fof(f1046,plain,(
  spl59_98 <=> type_(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_98])])).
fof(f36652,plain,(
  ~type_(xyz) | (~spl59_1 | ~spl59_2 | spl59_3 | ~spl59_749 | ~spl59_1701)),
  inference(subsumption_resolution,[],[f36651,f655])).
fof(f655,plain,(
  xyz != zyx | spl59_3),
  inference(avatar_component_clause,[],[f654])).
fof(f654,plain,(
  spl59_3 <=> xyz = zyx),
  introduced(avatar_definition,[new_symbols(naming,[spl59_3])])).
fof(f36651,plain,(
  xyz = zyx | ~type_(xyz) | (~spl59_1 | ~spl59_2 | ~spl59_749 | ~spl59_1701)),
  inference(subsumption_resolution,[],[f36650,f647])).
fof(f647,plain,(
  kind(zyx) | ~spl59_1),
  inference(avatar_component_clause,[],[f646])).
fof(f646,plain,(
  spl59_1 <=> kind(zyx)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1])])).
fof(f36650,plain,(
  ~kind(zyx) | xyz = zyx | ~type_(xyz) | (~spl59_2 | ~spl59_749 | ~spl59_1701)),
  inference(subsumption_resolution,[],[f36631,f651])).
fof(f651,plain,(
  kind(xyz) | ~spl59_2),
  inference(avatar_component_clause,[],[f650])).
fof(f650,plain,(
  spl59_2 <=> kind(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_2])])).
fof(f36631,plain,(
  ~kind(xyz) | ~kind(zyx) | xyz = zyx | ~type_(xyz) | (~spl59_749 | ~spl59_1701)),
  inference(resolution,[],[f36582,f5779])).
fof(f5779,plain,(
  ( ! [X3,X4,X5] : (~iof(sK46(X3),X4,X5) | ~kind(X3) | ~kind(X4) | X3 = X4 | ~type_(X3)) ) | ~spl59_749),
  inference(avatar_component_clause,[],[f5778])).
fof(f5778,plain,(
  spl59_749 <=> ! [X4,X5,X3] : (~iof(sK46(X3),X4,X5) | ~kind(X3) | ~kind(X4) | X3 = X4 | ~type_(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_749])])).
fof(f36582,plain,(
  iof(sK46(xyz),zyx,sK47(xyz)) | ~spl59_1701),
  inference(avatar_component_clause,[],[f36581])).
fof(f36581,plain,(
  spl59_1701 <=> iof(sK46(xyz),zyx,sK47(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1701])])).
fof(f36583,plain,(
  spl59_1701 | ~spl59_1647 | ~spl59_1666),
  inference(avatar_split_clause,[],[f28154,f28126,f27926,f36581])).
fof(f27926,plain,(
  spl59_1647 <=> iof(sK46(xyz),xyz,sK47(xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1647])])).
fof(f28126,plain,(
  spl59_1666 <=> ! [X0,X1] : (iof(X0,zyx,X1) | ~iof(X0,xyz,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1666])])).
fof(f28154,plain,(
  iof(sK46(xyz),zyx,sK47(xyz)) | (~spl59_1647 | ~spl59_1666)),
  inference(resolution,[],[f28127,f27927])).
fof(f27927,plain,(
  iof(sK46(xyz),xyz,sK47(xyz)) | ~spl59_1647),
  inference(avatar_component_clause,[],[f27926])).
fof(f28127,plain,(
  ( ! [X0,X1] : (~iof(X0,xyz,X1) | iof(X0,zyx,X1)) ) | ~spl59_1666),
  inference(avatar_component_clause,[],[f28126])).
fof(f28128,plain,(
  spl59_1666 | ~spl59_401 | ~spl59_1642),
  inference(avatar_split_clause,[],[f27827,f27823,f2912,f28126])).
fof(f2912,plain,(
  spl59_401 <=> ! [X0,X3,X2,X1] : (~iof(X3,X0,X2) | iof(X3,X1,X2) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_401])])).
fof(f27823,plain,(
  spl59_1642 <=> specializes(xyz,zyx)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1642])])).
fof(f27827,plain,(
  ( ! [X0,X1] : (iof(X0,zyx,X1) | ~iof(X0,xyz,X1)) ) | (~spl59_401 | ~spl59_1642)),
  inference(resolution,[],[f27824,f2913])).
fof(f2913,plain,(
  ( ! [X2,X3,X0,X1] : (~specializes(X0,X1) | iof(X3,X1,X2) | ~iof(X3,X0,X2)) ) | ~spl59_401),
  inference(avatar_component_clause,[],[f2912])).
fof(f27824,plain,(
  specializes(xyz,zyx) | ~spl59_1642),
  inference(avatar_component_clause,[],[f27823])).
fof(f27928,plain,(
  spl59_1647 | ~spl59_90 | ~spl59_1094 | ~spl59_1639),
  inference(avatar_split_clause,[],[f27808,f27745,f13700,f1010,f27926])).
fof(f1010,plain,(
  spl59_90 <=> sortal(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_90])])).
fof(f13700,plain,(
  spl59_1094 <=> ! [X1] : (iof(sK46(X1),sK21(X1),sK47(X1)) | ~sortal(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1094])])).
fof(f27745,plain,(
  spl59_1639 <=> xyz = sK21(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1639])])).
fof(f27808,plain,(
  iof(sK46(xyz),xyz,sK47(xyz)) | (~spl59_90 | ~spl59_1094 | ~spl59_1639)),
  inference(subsumption_resolution,[],[f27766,f1011])).
fof(f1011,plain,(
  sortal(xyz) | ~spl59_90),
  inference(avatar_component_clause,[],[f1010])).
fof(f27766,plain,(
  iof(sK46(xyz),xyz,sK47(xyz)) | ~sortal(xyz) | (~spl59_1094 | ~spl59_1639)),
  inference(superposition,[],[f13701,f27746])).
fof(f27746,plain,(
  xyz = sK21(xyz) | ~spl59_1639),
  inference(avatar_component_clause,[],[f27745])).
fof(f13701,plain,(
  ( ! [X1] : (iof(sK46(X1),sK21(X1),sK47(X1)) | ~sortal(X1)) ) | ~spl59_1094),
  inference(avatar_component_clause,[],[f13700])).
fof(f27825,plain,(
  spl59_1642 | ~spl59_111 | ~spl59_1640),
  inference(avatar_split_clause,[],[f27817,f27810,f1098,f27823])).
fof(f1098,plain,(
  spl59_111 <=> ! [X0] : (specializes(X0,zyx) | ~specializes(X0,xyz))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_111])])).
fof(f27810,plain,(
  spl59_1640 <=> specializes(xyz,xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1640])])).
fof(f27817,plain,(
  specializes(xyz,zyx) | (~spl59_111 | ~spl59_1640)),
  inference(resolution,[],[f27811,f1099])).
fof(f1099,plain,(
  ( ! [X0] : (~specializes(X0,xyz) | specializes(X0,zyx)) ) | ~spl59_111),
  inference(avatar_component_clause,[],[f1098])).
fof(f27811,plain,(
  specializes(xyz,xyz) | ~spl59_1640),
  inference(avatar_component_clause,[],[f27810])).
fof(f27812,plain,(
  spl59_1640 | ~spl59_90 | ~spl59_102 | ~spl59_1639),
  inference(avatar_split_clause,[],[f27807,f27745,f1062,f1010,f27810])).
fof(f1062,plain,(
  spl59_102 <=> ! [X0] : (specializes(X0,sK21(X0)) | ~sortal(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_102])])).
fof(f27807,plain,(
  specializes(xyz,xyz) | (~spl59_90 | ~spl59_102 | ~spl59_1639)),
  inference(subsumption_resolution,[],[f27749,f1011])).
fof(f27749,plain,(
  specializes(xyz,xyz) | ~sortal(xyz) | (~spl59_102 | ~spl59_1639)),
  inference(superposition,[],[f1063,f27746])).
fof(f1063,plain,(
  ( ! [X0] : (specializes(X0,sK21(X0)) | ~sortal(X0)) ) | ~spl59_102),
  inference(avatar_component_clause,[],[f1062])).
fof(f27747,plain,(
  spl59_1639 | ~spl59_2 | ~spl59_90 | ~spl59_1637),
  inference(avatar_split_clause,[],[f26823,f26760,f1010,f650,f27745])).
fof(f26760,plain,(
  spl59_1637 <=> ! [X9] : (~sortal(X9) | ~kind(X9) | sK21(X9) = X9)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_1637])])).
fof(f26823,plain,(
  xyz = sK21(xyz) | (~spl59_2 | ~spl59_90 | ~spl59_1637)),
  inference(subsumption_resolution,[],[f26768,f1011])).
fof(f26768,plain,(
  ~sortal(xyz) | xyz = sK21(xyz) | (~spl59_2 | ~spl59_1637)),
  inference(resolution,[],[f26761,f651])).
fof(f26761,plain,(
  ( ! [X9] : (~kind(X9) | ~sortal(X9) | sK21(X9) = X9) ) | ~spl59_1637),
  inference(avatar_component_clause,[],[f26760])).
fof(f26762,plain,(
  spl59_1637 | ~spl59_77 | ~spl59_160 | ~spl59_749 | ~spl59_1094),
  inference(avatar_split_clause,[],[f13731,f13700,f5778,f1299,f950,f26760])).
fof(f950,plain,(
  spl59_77 <=> ! [X0] : (kind(sK21(X0)) | ~sortal(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_77])])).
fof(f1299,plain,(
  spl59_160 <=> ! [X1] : (~sortal(X1) | type_(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_160])])).
fof(f13731,plain,(
  ( ! [X9] : (~sortal(X9) | ~kind(X9) | sK21(X9) = X9) ) | (~spl59_77 | ~spl59_160 | ~spl59_749 | ~spl59_1094)),
  inference(subsumption_resolution,[],[f13730,f951])).
fof(f951,plain,(
  ( ! [X0] : (kind(sK21(X0)) | ~sortal(X0)) ) | ~spl59_77),
  inference(avatar_component_clause,[],[f950])).
fof(f13730,plain,(
  ( ! [X9] : (~sortal(X9) | ~kind(X9) | ~kind(sK21(X9)) | sK21(X9) = X9) ) | (~spl59_160 | ~spl59_749 | ~spl59_1094)),
  inference(subsumption_resolution,[],[f13712,f1300])).
fof(f1300,plain,(
  ( ! [X1] : (~sortal(X1) | type_(X1)) ) | ~spl59_160),
  inference(avatar_component_clause,[],[f1299])).
fof(f13712,plain,(
  ( ! [X9] : (~sortal(X9) | ~kind(X9) | ~kind(sK21(X9)) | sK21(X9) = X9 | ~type_(X9)) ) | (~spl59_749 | ~spl59_1094)),
  inference(resolution,[],[f13701,f5779])).
fof(f13702,plain,(
  spl59_1094 | ~spl59_160 | ~spl59_173 | ~spl59_697),
  inference(avatar_split_clause,[],[f5359,f5340,f1376,f1299,f13700])).
fof(f1376,plain,(
  spl59_173 <=> ! [X0] : (iof(sK46(X0),X0,sK47(X0)) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_173])])).
fof(f5340,plain,(
  spl59_697 <=> ! [X2,X0,X1] : (iof(X0,sK21(X1),X2) | ~iof(X0,X1,X2) | ~sortal(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_697])])).
fof(f5359,plain,(
  ( ! [X1] : (iof(sK46(X1),sK21(X1),sK47(X1)) | ~sortal(X1)) ) | (~spl59_160 | ~spl59_173 | ~spl59_697)),
  inference(subsumption_resolution,[],[f5353,f1300])).
fof(f5353,plain,(
  ( ! [X1] : (iof(sK46(X1),sK21(X1),sK47(X1)) | ~sortal(X1) | ~type_(X1)) ) | (~spl59_173 | ~spl59_697)),
  inference(resolution,[],[f5341,f1377])).
fof(f1377,plain,(
  ( ! [X0] : (iof(sK46(X0),X0,sK47(X0)) | ~type_(X0)) ) | ~spl59_173),
  inference(avatar_component_clause,[],[f1376])).
fof(f5341,plain,(
  ( ! [X2,X0,X1] : (~iof(X0,X1,X2) | iof(X0,sK21(X1),X2) | ~sortal(X1)) ) | ~spl59_697),
  inference(avatar_component_clause,[],[f5340])).
fof(f5780,plain,(
  spl59_749 | ~spl59_173 | ~spl59_436),
  inference(avatar_split_clause,[],[f3316,f3183,f1376,f5778])).
fof(f3183,plain,(
  spl59_436 <=> ! [X4,X0,X3,X2,X1] : (~kind(X1) | ~iof(X0,X1,X2) | ~kind(X3) | ~iof(X0,X3,X4) | X1 = X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_436])])).
fof(f3316,plain,(
  ( ! [X3,X4,X5] : (~iof(sK46(X3),X4,X5) | ~kind(X3) | ~kind(X4) | X3 = X4 | ~type_(X3)) ) | (~spl59_173 | ~spl59_436)),
  inference(resolution,[],[f3184,f1377])).
fof(f3184,plain,(
  ( ! [X2,X3,X0,X1,X4] : (~iof(X0,X3,X4) | ~iof(X0,X1,X2) | ~kind(X3) | ~kind(X1) | X1 = X3) ) | ~spl59_436),
  inference(avatar_component_clause,[],[f3183])).
fof(f5342,plain,(
  spl59_697 | ~spl59_102 | ~spl59_401),
  inference(avatar_split_clause,[],[f2922,f2912,f1062,f5340])).
fof(f2922,plain,(
  ( ! [X2,X0,X1] : (iof(X0,sK21(X1),X2) | ~iof(X0,X1,X2) | ~sortal(X1)) ) | (~spl59_102 | ~spl59_401)),
  inference(resolution,[],[f2913,f1063])).
fof(f3185,plain,(
  spl59_436),
  inference(avatar_split_clause,[],[f644,f3183])).
fof(f644,plain,(
  ( ! [X2,X3,X0,X1,X4] : (~kind(X1) | ~iof(X0,X1,X2) | ~kind(X3) | ~iof(X0,X3,X4) | X1 = X3) )),
  inference(subsumption_resolution,[],[f633,f630])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((world(X2) & type_(X1)) | ~iof(X0,X1,X2))),
  inference(ennf_transformation,[],[f48])).
fof(f48,axiom,(
  ! [X0,X1,X2] : (iof(X0,X1,X2) => (world(X2) & type_(X1)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_dIof)).
fof(f633,plain,(
  ( ! [X2,X3,X0,X1,X4] : (~world(X2) | ~kind(X1) | ~iof(X0,X1,X2) | ~kind(X3) | ~iof(X0,X3,X4) | X1 = X3) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : (X1 = X3 | ~iof(X0,X3,X4) | ~kind(X3)) | ~iof(X0,X1,X2) | ~kind(X1) | ~world(X2))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : (X1 = X3 | ~iof(X0,X3,X4) | ~kind(X3)) | (~iof(X0,X1,X2) | ~kind(X1) | ~world(X2)))),
  inference(ennf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((iof(X0,X1,X2) & kind(X1) & world(X2)) => ~? [X3,X4] : (X1 != X3 & iof(X0,X3,X4) & kind(X3)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6,X11,X2] : ((iof(X6,X11,X2) & kind(X11) & world(X2)) => ~? [X12,X10] : (X11 != X12 & iof(X6,X12,X10) & kind(X12)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_uniqueKind_a22)).
fof(f2914,plain,(
  spl59_401 | ~spl59_129 | ~spl59_398),
  inference(avatar_split_clause,[],[f2846,f2843,f1170,f2912])).
fof(f1170,plain,(
  spl59_129 <=> ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_129])])).
fof(f2843,plain,(
  spl59_398 <=> ! [X0,X3,X2,X1] : (~world(X2) | ~iof(X3,X0,X2) | iof(X3,X1,X2) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_398])])).
fof(f2846,plain,(
  ( ! [X2,X3,X0,X1] : (~iof(X3,X0,X2) | iof(X3,X1,X2) | ~specializes(X0,X1)) ) | (~spl59_129 | ~spl59_398)),
  inference(subsumption_resolution,[],[f2844,f1171])).
fof(f1171,plain,(
  ( ! [X2,X0,X1] : (~iof(X0,X1,X2) | world(X2)) ) | ~spl59_129),
  inference(avatar_component_clause,[],[f1170])).
fof(f2844,plain,(
  ( ! [X2,X3,X0,X1] : (~world(X2) | ~iof(X3,X0,X2) | iof(X3,X1,X2) | ~specializes(X0,X1)) ) | ~spl59_398),
  inference(avatar_component_clause,[],[f2843])).
fof(f2845,plain,(
  spl59_398),
  inference(avatar_split_clause,[],[f608,f2843])).
fof(f608,plain,(
  ( ! [X2,X3,X0,X1] : (~world(X2) | ~iof(X3,X0,X2) | iof(X3,X1,X2) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (! [X3] : (iof(X3,X1,X2) | ~iof(X3,X0,X2)) | ~world(X2)) & type_(X1) & type_(X0)))),
  inference(ennf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (world(X2) => ! [X3] : (iof(X3,X0,X2) => iof(X3,X1,X2))) & type_(X1) & type_(X0)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X4,X5] : (specializes(X4,X5) <=> (! [X2] : (world(X2) => ! [X6] : (iof(X6,X4,X2) => iof(X6,X5,X2))) & type_(X5) & type_(X4)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_specialization_a5)).
fof(f1378,plain,(
  spl59_173),
  inference(avatar_split_clause,[],[f504,f1376])).
fof(f504,plain,(
  ( ! [X0] : (iof(sK46(X0),X0,sK47(X0)) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1,X2] : iof(X1,X0,X2))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_dType_a1)).
fof(f1301,plain,(
  spl59_160 | ~spl59_81 | ~spl59_102),
  inference(avatar_split_clause,[],[f1297,f1062,f966,f1299])).
fof(f966,plain,(
  spl59_81 <=> ! [X0,X1] : (~specializes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_81])])).
fof(f1297,plain,(
  ( ! [X1] : (~sortal(X1) | type_(X1)) ) | (~spl59_81 | ~spl59_102)),
  inference(resolution,[],[f1063,f967])).
fof(f967,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl59_81),
  inference(avatar_component_clause,[],[f966])).
fof(f1172,plain,(
  spl59_129),
  inference(avatar_split_clause,[],[f630,f1170])).
fof(f1100,plain,(
  spl59_111),
  inference(avatar_split_clause,[],[f506,f1098])).
fof(f506,plain,(
  ( ! [X0] : (specializes(X0,zyx) | ~specializes(X0,xyz)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : (specializes(X0,xyz) <=> specializes(X0,zyx))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X31] : (specializes(X31,xyz) <=> specializes(X31,zyx))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',axiom1710)).
fof(f1064,plain,(
  spl59_102),
  inference(avatar_split_clause,[],[f413,f1062])).
fof(f413,plain,(
  ( ! [X0] : (specializes(X0,sK21(X0)) | ~sortal(X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : (sortal(X0) <=> (? [X1] : (specializes(X0,X1) & kind(X1)) & endurantType(X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X13] : (sortal(X13) <=> (? [X11] : (specializes(X13,X11) & kind(X11)) & endurantType(X13)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_dSortal_a23)).
fof(f1048,plain,(
  spl59_98 | ~spl59_30 | ~spl59_95),
  inference(avatar_split_clause,[],[f1036,f1033,f762,f1046])).
fof(f762,plain,(
  spl59_30 <=> ! [X0] : (~endurantType(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_30])])).
fof(f1033,plain,(
  spl59_95 <=> endurantType(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_95])])).
fof(f1036,plain,(
  type_(xyz) | (~spl59_30 | ~spl59_95)),
  inference(resolution,[],[f1034,f763])).
fof(f763,plain,(
  ( ! [X0] : (~endurantType(X0) | type_(X0)) ) | ~spl59_30),
  inference(avatar_component_clause,[],[f762])).
fof(f1034,plain,(
  endurantType(xyz) | ~spl59_95),
  inference(avatar_component_clause,[],[f1033])).
fof(f1035,plain,(
  spl59_95 | ~spl59_43 | ~spl59_90),
  inference(avatar_split_clause,[],[f1013,f1010,f814,f1033])).
fof(f814,plain,(
  spl59_43 <=> ! [X0] : (endurantType(X0) | ~sortal(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_43])])).
fof(f1013,plain,(
  endurantType(xyz) | (~spl59_43 | ~spl59_90)),
  inference(resolution,[],[f1011,f815])).
fof(f815,plain,(
  ( ! [X0] : (~sortal(X0) | endurantType(X0)) ) | ~spl59_43),
  inference(avatar_component_clause,[],[f814])).
fof(f1012,plain,(
  spl59_90 | ~spl59_38 | ~spl59_86),
  inference(avatar_split_clause,[],[f993,f990,f794,f1010])).
fof(f794,plain,(
  spl59_38 <=> ! [X0] : (sortal(X0) | ~rigidSortal(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_38])])).
fof(f990,plain,(
  spl59_86 <=> rigidSortal(xyz)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_86])])).
fof(f993,plain,(
  sortal(xyz) | (~spl59_38 | ~spl59_86)),
  inference(resolution,[],[f991,f795])).
fof(f795,plain,(
  ( ! [X0] : (~rigidSortal(X0) | sortal(X0)) ) | ~spl59_38),
  inference(avatar_component_clause,[],[f794])).
fof(f991,plain,(
  rigidSortal(xyz) | ~spl59_86),
  inference(avatar_component_clause,[],[f990])).
fof(f992,plain,(
  spl59_86 | ~spl59_2 | ~spl59_72),
  inference(avatar_split_clause,[],[f982,f930,f650,f990])).
fof(f930,plain,(
  spl59_72 <=> ! [X0] : (rigidSortal(X0) | ~kind(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl59_72])])).
fof(f982,plain,(
  rigidSortal(xyz) | (~spl59_2 | ~spl59_72)),
  inference(resolution,[],[f931,f651])).
fof(f931,plain,(
  ( ! [X0] : (~kind(X0) | rigidSortal(X0)) ) | ~spl59_72),
  inference(avatar_component_clause,[],[f930])).
fof(f968,plain,(
  spl59_81),
  inference(avatar_split_clause,[],[f568,f966])).
fof(f568,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) | ~specializes(X0,X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1] : (specializes(X0,X1) => (type_(X1) & type_(X0)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_dSpecializes)).
fof(f952,plain,(
  spl59_77),
  inference(avatar_split_clause,[],[f412,f950])).
fof(f412,plain,(
  ( ! [X0] : (kind(sK21(X0)) | ~sortal(X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f932,plain,(
  spl59_72),
  inference(avatar_split_clause,[],[f555,f930])).
fof(f555,plain,(
  ( ! [X0] : (rigidSortal(X0) | ~kind(X0)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,axiom,(
  ! [X0] : ((subkind(X0) | kind(X0)) <=> rigidSortal(X0))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_rigidSortal_taxonomy)).
fof(f816,plain,(
  spl59_43),
  inference(avatar_split_clause,[],[f414,f814])).
fof(f414,plain,(
  ( ! [X0] : (endurantType(X0) | ~sortal(X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f796,plain,(
  spl59_38),
  inference(avatar_split_clause,[],[f381,f794])).
fof(f381,plain,(
  ( ! [X0] : (sortal(X0) | ~rigidSortal(X0)) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : (rigidSortal(X0) <=> (sortal(X0) & rigid(X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X8] : (rigidSortal(X8) <=> (sortal(X8) & rigid(X8)))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_rigidSortalsAreRigidAndSortal_xx)).
fof(f764,plain,(
  spl59_30),
  inference(avatar_split_clause,[],[f359,f762])).
fof(f359,plain,(
  ( ! [X0] : (~endurantType(X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : (type_(X0) | (~perdurantType(X0) & ~endurantType(X0)))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0] : ((perdurantType(X0) | endurantType(X0)) => type_(X0))),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',ax_type_taxonomy)).
fof(f656,plain,(
  ~spl59_3),
  inference(avatar_split_clause,[],[f312,f654])).
fof(f312,plain,(
  xyz != zyx),
  inference(cnf_transformation,[],[f135])).
fof(f135,axiom,(
  xyz != zyx),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',diff1709)).
fof(f652,plain,(
  spl59_2),
  inference(avatar_split_clause,[],[f311,f650])).
fof(f311,plain,(
  kind(xyz)),
  inference(cnf_transformation,[],[f132])).
fof(f132,axiom,(
  kind(xyz)),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',axiom1708)).
fof(f648,plain,(
  spl59_1),
  inference(avatar_split_clause,[],[f310,f646])).
fof(f310,plain,(
  kind(zyx)),
  inference(cnf_transformation,[],[f133])).
fof(f133,axiom,(
  kind(zyx)),
  file('outputs/ufo/cis/ci_ufo_kind_specializes_@_2.in',axiom1709)).
% SZS output end Proof for ci_ufo_kind_specializes_@_2
% ------------------------------
% Version: Vampire 4.7 (commit af7d05f97 on 2022-09-30 11:01:41 +0200)
% Termination reason: Refutation

% Memory used [KB]: 31982
% Time elapsed: 0.708 s
% ------------------------------
% ------------------------------
% Success in time 0.717 s
